Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

### Exercise: Heart Disease / Scikit-Learn

Version: 3.6, Summer Semester 2024

Follow the steps to load the heart disease database and to train & score various classifiers with it.
The database is a slightly modified version from the [UCI Machine Learning Repository](https://archive.ics.uci.edu/ml/datasets/heart+Disease).

## Your Name

Replace the `raise NotImplementedError` with the code `myname = ""` and assign your name to the variable:

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: name-test, points: 1, grade: true, solution: false -->
```python
assert myname != "", "myname should not be empty"
```

## 1: Read about the Dataset

Check out the dataset description of [UCI](https://archive.ics.uci.edu/ml/datasets/heart+Disease). Our dataset is the Cleveland set and uses the smaller variant with 14 attributes. Take a look at column 3 ("chest pain" / "cp"). The values are either 1, 2, 3 or 4 in the dataset. 

Change the code below so that the Python dictionary contains the names of the four different types of chest pain that the dataset distinguishes.

*Hint: you'll find the answer in the "Attribute information" section of the linked web page.*

<!-- grade_id: chest-pain, points: 0, grade: false, solution: true -->
```python
# Replace xxx with the textual names of the different types of chest pain
pain_types = {1: "xxx",
             2: "xxx",
             3: "xxx",
             4: "xxx"}
# When you have finished replacing the values, delete the exception in the line
# below and run the next block with the asserts to see if you got it right.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: chest-pain-test, points: 3, grade: true, solution: false -->
```python
import hashlib
assert hashlib.blake2b(pain_types[1].casefold().encode()).hexdigest() == '10de5546a056feb6522445a2e94492fa20419030d70694bd97fc67016906936dd8a8791e286becb0bb5f8087986bd2621f2f06d639523f653f955a2d45e78d6f'
assert hashlib.blake2b(pain_types[2].casefold().encode()).hexdigest() == '884a2e513fec27229c985dc4b01d4935d42e90e6c71d1f5d519b38e7369ae23be525d803dadf9e102ddb91beb6feef4a4d71d8a9636c1a4106f87c1aee1c2640'
assert hashlib.blake2b(pain_types[3].casefold().encode()).hexdigest() == 'cb8436129905c921a1930a01e2ae2ebe74d5562826802d6ea27f183c6703a3a58d9e8cc94f5eb53e48928abcefdd88e8ac50e153c99f0870005e327fcd4d1c89'
assert hashlib.blake2b(pain_types[4].casefold().encode()).hexdigest() == '56f3734df8320938c6c2e66fec6ea97f04e6f5f7e5cc27c57da98fc99ba277442240d7c0a9d4aa12816199274de33843c591d6d61037cd5a12f347f50850bd5b'
```

## 2: Imports

Use the following block for all the imports you need for your notebook.
This will include `pandas` (as `pd`), `numpy` (as `np`), `seaborn` (as `sns`) and all the parts of `sklearn` you will need.
Also configure `matplotlib` to draw graphics inline. If during this notebook you discover that you need additional imports, just come back to this cell, add the import here and execute the cell again. It's often easier to have all imports in one place.

<!-- grade_id: imports, points: 0, grade: false, solution: true -->
```python
# Place all neccessary imports here
# YOUR CODE HERE
raise NotImplementedError()
```

## 3: Load the CSV

Load the `heart_disease-fhstp-withmissing.csv` file into a Pandas dataframe variable called `df_heart` and print its head to check if importing worked. The URL for the file is: `https://raw.githubusercontent.com/andijakl/MachineLearning/main/lab%202%20-%20scikit%20learn%20-%20heart%20disease/heart_disease-fhstp-withmissing.csv`

Our dataset contains some missing values. If you take a look at the file, you will see that some values are `?`. By default, these are not recognized as missing values by the CSV reader; instead, the whole column would be interpreted as string value, so that the `?` character is stored as text. To improve the situation, specify `'?'` for the `na_values` parameter of the `read_csv` function. This will correctly interpret `?` as missing data.

<!-- grade_id: read-csv, points: 0, grade: false, solution: true -->
```python
file_url = 'https://raw.githubusercontent.com/andijakl/MachineLearning/main/lab%202%20-%20scikit%20learn%20-%20heart%20disease/heart_disease-fhstp-withmissing.csv'
# Load the dataset into a variable called df_heart
# Specify that ? should be recognized as na value
# Print the head of the dataframe to see if importing worked correctly.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: read-csv-test, points: 2, grade: true, solution: false -->
```python
# Tests if data has been loaded
assert len(df_heart) == 303, "Imported data should have 303 samples"
```

Next, let's see how many missing values we have and in which columns these are. From the dataframe, run the `isna()` function. This will result in a new dataframe containing true / false values whether the value is na or not. This doesn't help us much, so add `.sum()` to get the sum of na values for each column. Store the resulting data in a variable called `sum_na`.

<!-- grade_id: sum-na, points: 0, grade: false, solution: true -->
```python
# Store the sum of na values for each column in a variable called sum_na.
# Print the variable
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: sum-na-test, points: 2, grade: true, solution: false -->
```python
assert sum_na['major vessels colored'] == 4
assert sum_na['thal'] == 2
```

As you can see, it affects few samples overall. We decide that we simply drop these rows from our dataset instead of imputing the missing values based on the other samples. For this, run `dropna()` on the dataset. To directly modify the source, use the `inplace=True` parameter.

<!-- grade_id: drop-na, points: 0, grade: false, solution: true -->
```python
# Drop missing values using the inplace version
# YOUR CODE HERE
raise NotImplementedError()
```

Now check how many samples we still have left using the `len()` function, supplying the dataframe as parameter. Save the new length in a variable called `df_heart_len` and print it.

<!-- grade_id: len-dropped, points: 0, grade: false, solution: true -->
```python
# Save the length of the dataset to a new variable df_heart_len and print it
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: len-dropped-test, points: 2, grade: true, solution: false -->
```python
assert df_heart_len == 297
assert len(df_heart) == 297
```

## 4: Describe

Use the describe function of the Pandas DataFrame to find out the count, mean, standard deviation & more from the dataset for each column.

<!-- grade_id: describe, points: 0, grade: false, solution: true -->
```python
# Call method to describe the dataset
# YOUR CODE HERE
raise NotImplementedError()
```

Based on the printed information you see about the dataset, answer a few questions and assign the numbers to the corresponding variables. Round your answers (up/down) to the next integer.

<!-- grade_id: dataset_data, points: 0, grade: false, solution: true -->
```python
# What is the mean age?
mean_age = 0;
# What is the maximum heart rate?
max_hr = 0;
# What is the standard deviation of the cholesterol level?
std_cholesterol = 0;
# Simply assign the values to the variables above and when finished,
# remove the line throwing the not implemented error.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dataset_data_tests, points: 2, grade: true, solution: false -->
```python
assert mean_age > 0
assert max_hr > 0
assert std_cholesterol > 0
```

## 5: Age Histogram

Print the histogram of the "age" column. Use `10` bins. This gives you a good understanding of the patients that were tested for heart diseases.

<!-- grade_id: histogram, points: 0, grade: false, solution: true -->
```python
# Call method to print the age histogram with 10 bins
# YOUR CODE HERE
raise NotImplementedError()
```

According to the age bins above (middle number), which age is most frequent in the dataset? Look at the histogram and assign what you think should be approximately in the middle of the highest bin to the following variable:

<!-- grade_id: histogram-question, points: 0, grade: false, solution: true -->
```python
# Change the variable to the approximate middle of the highest bin in the histogram above.
# This answers: from what age are most people in the dataset? Assign that age to most_age.
most_age = -1
# Again, assign your answer to the variable above and then delete the exception
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: histogram-test, points: 2, grade: true, solution: false -->
```python
assert most_age != -1
```

## 6: Violin Plot

Use the Violin Plot of Seaborn (add the import statement at the beginning of the file!) to print the target class (`"diameter narrowing"` -> *diagnosis of heart disease (angiographic disease status), Value 0: < 50% diameter narrowing*) on the `x` axis and `"max HR"` on the `y` axis.

<!-- grade_id: violin-plot, points: 0, grade: false, solution: true -->
```python
# Plot a violin plot with the parameters stated above
# YOUR CODE HERE
raise NotImplementedError()
```

Based on the violin plot, answer: is the max HR higher if the target class is 0 or 1?

<!-- grade_id: max-hr-question, points: 0, grade: false, solution: true -->
```python
# Change the variable max_hr_higher to 
# 0 if the maximum heart rate is higher if there is no heart disease;
# 1 if the maximum heart rate is higher if there is heart disease
max_hr_higher = -1
# Assign the value to the variable above and delete the exception
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: max-hr-test, points: 2, grade: true, solution: false -->
```python
assert max_hr_higher != -1
```

## 7: Prepare Data for Classification

Our CSV contains all the data plus the target class values. Remove (`pop`) the target class into an extra variable (`y`).
Then split the dataframe into the 4 variables for training & test data. Use the random state `10`. You don't need to specify the size of the training and testing set; according to the [train_test_split documentation](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html), the default size of the test data will be 0.25, the training data 0.75.

<!-- grade_id: target-class, points: 0, grade: false, solution: true -->
```python
# Pop the target class from df_heart, convert it to a numpy array stored in a variable y
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: target-class-test, points: 3, grade: true, solution: false -->
```python
assert y.shape == (297,)
assert y[1] == 1
assert df_heart.shape == (297,13)
```

<!-- grade_id: split, points: 0, grade: false, solution: true -->
```python
# Split the dataframe and the target classes using the random state 10 
# into the variables X_train, X_test, y_train and y_test
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: split-test, points: 3, grade: true, solution: false -->
```python
assert X_train.shape == (222,13)
assert y_train.shape == (222,)
assert X_test.iloc[0]['age'] == 60.0
assert X_train.iloc[5]['age'] == 46.0
```

## 8: Decision Tree

Create a decision tree in a new variable `clf_dtc`, fit it with the training data, visualize the trained decision tree, and score the classifier with the test data (store the score in a variable called `clf_dtc_score`).

Configure the classifier to use a min sample split of `20` and a random state `10`. See the [documentation](https://scikit-learn.org/stable/modules/generated/sklearn.tree.DecisionTreeClassifier.html) for how the corresponding parameters are called.

<!-- grade_id: dtc, points: 0, grade: false, solution: true -->
```python
# Store the classifier in a variable called clf_dtc and create it with the configuration stated above
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dtc-fit, points: 0, grade: false, solution: true -->
```python
# Fit the classifier
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dtc-score, points: 0, grade: false, solution: true -->
```python
# Score the classifier and store the value in clf_dtc_score
# Also print the resulting score
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dtc-visualize, points: 2, grade: true, solution: true -->
```python
# Plot the decision tree
# Optional: to get a better resolution and to easily re-use the image 
# (e.g., in a Master Thesis), also save the figure: 
# Add the following import to your import block above:
# import matplotlib.pyplot as plt
# Then use: plt.savefig("filename.svg") to save the tree as vector graphic.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dtc-test, points: 2, grade: true, solution: false -->
```python
assert clf_dtc.min_samples_split == 20
assert clf_dtc_score > 0
```

## 9: Random Forest

Test a random forest classifier (`RandomForestClassifier`) to see if it performs better than the single decision tree. Use a max depth of `3`, `20` estimators and a random state of `10`.
Again, fit it with the train data and test with the test data.

<!-- grade_id: rf, points: 0, grade: false, solution: true -->
```python
# Store the classifier in a variable called clf_randomforest
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: rf-fit-score, points: 0, grade: false, solution: true -->
```python
# Fit and score the classifier. Use the correct sets for training / testing.
# Store the scoring result in the variable clf_randomforest_score.
# Also print the score results variable for reference.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: rf-test, points: 2, grade: true, solution: false -->
```python
assert clf_randomforest.random_state == 10
assert clf_randomforest_score > 0
```

## 10: Prediction

Use the random forest classifier to predict the class of an unknown patient. Use the following measured features:
`[[23.0,0.0,3.0,120.0,233.0,1.0,2.0,173.0,1.0,0.6,2.0,1.0,6.0]]`

Store the results in a variable called `pred`. Also print the prediction for your reference.

*Note:* you might get a warning stating that our supplied data doesn't have valid feature names, even though the model was fitted (trained) with feature names. We supplied the whole dataframe for training, including the header row that contained the feature names. Therefore, Scikit-Learn is now warning us that we need to be careful to have the same order of values, as we're now supplying raw data without any column labels. You can ignore this warning here. Possible solutions would be to add the feature labels to a dataframe for prediction as well, or training the classifier just with the values as a numpy array instead of the full dataframe.

<!-- grade_id: predict, points: 0, grade: false, solution: true -->
```python
# Code to precict for the measurements listed above and store the result in the variable pred
# Note: you can ignore the warning if you supply just the raw data and not a dataframe with feature names like the classifier was trained on.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: predict-test, points: 1, grade: true, solution: false -->
```python
assert clf_randomforest.predict([[23.0,0.0,3.0,120.0,233.0,1.0,2.0,173.0,1.0,0.6,2.0,1.0,6.0]]) == [1], "Result should be 1"
```

Does the patient have heart disease? Assign `True` or `False` to the `heart_disease` variable, based on the prediction generated by the classifier.

<!-- grade_id: heart_disease-question, points: 0, grade: false, solution: true -->
```python
# Does the patient have heart disease? -> True / False?
heart_disease = -1
# Assign the value to the variable above and remove the exception below
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: heart_disease-test, points: 2, grade: true, solution: false -->
```python
assert (heart_disease == True or heart_disease == False)
```

## 11: Scaling

Many classifiers like k-NN, Support Vector Machines and Neural Networks work best with scaled features. Therefore, we need to scale these to make the features more comparable. For example, k-NN computes the euclidian distance - as such, normalized / standardized feature values are important.

Create the `StandardScaler`. Then use `fit_transform()` to analyze the training data and to scale it in one step. Afterwards, use `transform()` to scale the test data (this applies the same scaling that was learned from the training data).
Make sure to save the scaled test + train data into new variables, e.g., `X_train_scaled` and `X_test_scaled`.

<!-- grade_id: scaler-create, points: 0, grade: false, solution: true -->
```python
# Create the standard scaler and store it in a variable called scaler
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: scaler-transform, points: 0, grade: false, solution: true -->
```python
# Correctly scale the train and test data into X_train_scaled and X_test_scaled
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: scaler-test, points: 4, grade: true, solution: false -->
```python
assert X_train_scaled.shape == (222,13)
assert X_test_scaled.shape == (75,13)
assert np.isclose(X_train_scaled[0][0], 1.0153, rtol=0.1), "Should be close to 1.1053"
```

## 12: k-NN Classification

Create the `KNeighborsClassifier` and call the variable `clf_knn`. Use the default settings (which is 5 neighbors). Next, fit the classifier on the *scaled* training data and score it with the *scaled* test data. Store the result in a variable `clf_knn_score`.

<!-- grade_id: knn, points: 0, grade: false, solution: true -->
```python
# Store the classifier in a variable clf_knn.
# Save the scoring result in a variable clf_knn_score.
# Print the score for your reference.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: knn-test, points: 2, grade: true, solution: false -->
```python
assert clf_knn.n_neighbors == 5
assert clf_knn_score > 0
```

## 13: Support Vector Machine

Create a support vector classifier for variable `clf_svc` (Hint: add to imports block at the beginning of the notebook: `from sklearn import svm`. Create the classifier in the following block using `svm.SVC`).

Set `gamma` to `"scale"` and a random state of `10`.

Next, fit the classifier on the *scaled* training data and score it with the *scaled* test data. Save the score in a variable `clf_svc_score`).

<!-- grade_id: svc-create, points: 0, grade: false, solution: true -->
```python
# Store the classifier in a variable called clf_svc
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: svc-fit, points: 0, grade: false, solution: true -->
```python
# Fit the classifier to the scaled training data
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: svc-score, points: 0, grade: false, solution: true -->
```python
# Score the classifier and assign the score to a variable called clf_svc_score.
# Print the score for your reference.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: svc-test, points: 2, grade: true, solution: false -->
```python
assert clf_svc.random_state == 10
assert clf_svc_score > 0
```

## 14: Neural Network

As a final classifier, create a neural network (multi-layer perceptron; add to the import block above: `from sklearn.neural_network import MLPClassifier`). As before, specify the random_state of `10`. Use `2000` maximum iterations when configuring the classifier (`max_iter`). Again, fit and score the classifier on the *scaled* data. Store the scoring result in `clf_mlp_score`.

<!-- grade_id: mlp-init, points: 0, grade: false, solution: true -->
```python
# Store classifier in variable called clf_mlp
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: mlp-fit, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: mlp-score, points: 0, grade: false, solution: true -->
```python
# Score the classifier and assign the score to a variable called clf_mlp_score
# Also print the score for your reference.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: mlp-test, points: 2, grade: true, solution: false -->
```python
assert clf_mlp_score > 0
assert clf_mlp.random_state == 10
```

## 15: Final Checks & the best classifier

To make sure that your whole Jupyter Notebook executes without issues, choose *Kernel/Runtime -> Restart Kernel and Run All Cells...*. Make sure it actually restarts executing all cells - if not, select the command again. Make sure all the lines you wrote execute and that you get the scores.

Answer in the following code block: which classifier got the highest accuracy for the data? Store your answer in a variable called `best_classifier` according to:



1.   Decision Tree
2.   Random Forest
3.   k-Nearest Neighbor
4.   Support Vector Machine
5.   Multi-Layer Perceptron (Neural Network)



<!-- grade_id: best-classifier, points: 0, grade: false, solution: true -->
```python
# Change the variable best_classifier to a number according
# to the reference above.
best_classifier = -1
# Once you've updated the variable, delete the exception line below
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: best-classifier-test, points: 3, grade: true, solution: false -->
```python
assert best_classifier != -1
```