Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Exercise 1: Development Basics

Version 5.2, Summer Semester 2024

## Part 2: Numpy

## Your Name

Replace the `raise NotImplementedError` with the code `myname = ""` and assign your name to the variable:

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: name-test, points: 1, grade: true, solution: false -->
```python
assert myname != "", "myname should not be empty"
```

# Numpy

A good overview to repeat from our class the concepts is in the [Data Science Foundations: Python Scientific Stack](https://www.linkedin.com/learning/data-science-foundations-python-scientific-stack/overview-numpy) online course on the LinkedIn Learning Platform. Watch chapter 3 - NumPy Basics (duration approx. 27min).

To start with numpy we need to import the numpy package. Execute the following lines also in your notebook by pressing shift+enter in the following code line:

```python
import numpy as np
np.random.seed(0) # we also reset the random number generator, so that everyone of us gets the same random numbers in this exercise
```

## Exercise 2.1: NumPy Arrays

Create following **numpy array** `arr1`: 

$$ arr1 = 
 \begin{pmatrix}
  1 & 0 & 1 & 0 & 1 & 0 & 1 & 0 & 1 & 0
 \end{pmatrix} $$

Create also a numpy array `arr2` with 10 *random* elements in the range of 0..1.

Execute the following lines with your own code:
* print the values of the elements of both arrays.
* add the two arrays to each other (using `+`, not using concatenation). Store the result in variable `arr_add`.
* multiply the two arrays elementwise. Store the result in variable `arr_mul`.
* calculate the dot product of the two arrays. Store the result in variable `arr_dot`.

<!-- grade_id: arr1, points: 0, grade: false, solution: true -->
```python
# Define arr1 and print it
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: arr2, points: 0, grade: false, solution: true -->
```python
# Define arr2 and print it
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: arr-add, points: 0, grade: false, solution: true -->
```python
# Calculate the sum of arr1 and arr2, store in arr_add
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: arr-mul, points: 0, grade: false, solution: true -->
```python
# Multiply arr1 with arr2, store in arr_mul
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: arr-dot, points: 0, grade: false, solution: true -->
```python
# Dot multiplication of arr1 and arr2, store in arr_dot
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: arr-test, points: 5, grade: true, solution: false -->
```python
assert arr1.shape == (10,)
assert arr2.shape == (10,)
assert arr_mul[1] == 0
assert arr_dot.size == 1
assert np.isclose(arr_dot, 2.976, rtol=0.1)
```

## Exercise 2.2: Matrix Access & Shape

Create following numpy 2d array in a variable `A`: 
$$ A = 
 \begin{pmatrix}
  2 & 4 & 2 & 6 & 3 \\
  1 & 3 & 2 & 8 & -3 \\
  9 & 3 & 5 & 7 & 2 \\
  -4 & 1 & 20 & 0 & -20
 \end{pmatrix} $$

<!-- grade_id: matrix-def, points: 0, grade: false, solution: true -->
```python
# Define matrix A
# YOUR CODE HERE
raise NotImplementedError()
```

What is the shape of the matrix you generated? Use `shape` based on `A` and store the result in a new variable called `mat_shape`. Also print the variable. Note that `shape` isn't a function, so you do not need `()` at the end.

<!-- grade_id: mat-shape, points: 0, grade: false, solution: true -->
```python
# Store the shape of A into a variable mat_shape
# (use the corresponding function and do not assign the values manually)
# YOUR CODE HERE
raise NotImplementedError()
```

How many rows does this matrix have? Access the information from variable `mat_shape` and store the value in a new variable `mat_rows`

<!-- grade_id: mat-rows, points: 0, grade: false, solution: true -->
```python
# Store the number of rows of matrix A in a new variable mat_rows (based on mat_shape)
# YOUR CODE HERE
raise NotImplementedError()
```

From Matrix A, select all rows *starting with* the row 1 (note: counting starts with 0), with all columns *until* the column 1. Store the result in a new variable `mat_slice`.

<!-- grade_id: mat_slice, points: 0, grade: false, solution: true -->
```python
# Create a variable mat_slice and store a slice according to the instructions above
# YOUR CODE HERE
raise NotImplementedError()
```

From Matrix A, access the row 2 (note: counting starts with 0). From these rows, select all columns. Store the slice in a new variable called `mat_single`. If each row is data from a patient and each column contains a measurement, this would equal getting the data of a single patient.

<!-- grade_id: mat_single, points: 0, grade: false, solution: true -->
```python
# Create a variable mat_single for row 2 and all columns
# YOUR CODE HERE
raise NotImplementedError()
```

From Matrix A, slice columns 2 and 0 (in this order) for all rows. Store the result in a variable `mat_columns`. If each row is data from a patient and each column contains a measurement, this would equal only looking at certain measurements for all patients.

<!-- grade_id: mat_columns, points: 0, grade: false, solution: true -->
```python
# Create a variable mat_columns containing columns 2 and 0 (in this order) for all rows.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: mat-test, points: 6, grade: true, solution: false -->
```python
assert A.shape == (4,5)
assert A[0,1] == 4
assert A[2,3] == 7
assert mat_shape == (4,5)
assert mat_rows > 0
assert mat_rows < 6
assert mat_slice.shape == (3,1)
assert mat_single.shape == (5,)
assert mat_columns.shape == (4,2)
```