# Exercise 1: Development Basics

Version 5.0, Summer Semester 2022

## Part 1: Python

If you want to learn more about Python, the following course is recommended: [Python Essential Training](https://www.linkedin.com/learning/python-essential-training-2) on LinkedIn Learning (or Lynda). Chapters 2-7 are mostly relevant.

There are also some cheat sheets you can use to remind yourself of the basic concepts in python:
* [Python 3 Cheat Sheet](https://groklearning-cdn.com/resources/cheatsheet-python-1.pdf) by GROK Learning
* [Python 3 Cheat Sheet](https://perso.limsi.fr/pointal/_media/python:cours:mementopython3-english.pdf) by Laurent Pointal

## Exercise 1.1: Your Name & Comments

Write your first name, family name and student ID number as Python comments:

<!-- grade_id: comment, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-1-15b94d1fa268> in <module>
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

## Exercise 1.2: String Variables and Assignments

Assign to a variable `s1` your first name as a string and to a variable `s2` your family name as a string. Concatinate those two strings into a new variable `s3` so that there is a space between the names, e.g. "Max Müller". Print variable `s3`. 

<!-- grade_id: string, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: string-test, points: 2, grade: true, solution: false -->
```python
assert len(s1) > 0
assert len(s2) > 0
```

## Exercise 1.3: Numeric Variables

Assign to a variable `x` the value 10 and to a variable `y` the value 5. You should perform basic arithmetic operations: a) add, b) subtract, c) divide, and d) multiply the two variables and save the results into new variables with the corresponding names (called `add`, `sub`, `div`, `mul`).

Output the results by calling the `print( )` function only once: concatenate all results to one string.

**Hint:** you can use the `f` syntax or the `format( )` function inside the print function to insert values of variables in a string. 

<!-- grade_id: variables, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: variables-test, points: 2, grade: true, solution: false -->
```python
assert x == 10
assert y == 5
assert add == 15
assert sub > 0
assert div > 0
assert mul > 0
```

**Additionally:** print the type of variable `x` and the type of the variable containing the result of the division (`div`).

<!-- grade_id: types, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Explain:** Why are the data types different? Write your explanation in a markdown cell below.

<!-- grade_id: types-reason, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 1.4: Functions

Define a function called `myMultiplication`, which should have two parameters `a` (default value: 5) and `b` (default value: 2). In the function you should implement an algorithm which calculates the multiplication of the two parameters by employing only one `for` or `while` loop and addition `+`. Do not use the multiplication operator `*`, instead use a loop to compute `a*b` by adding up `a` to `a` for `b` times). 

a) Call the function once without parameters and print the result. This should use the standard values for `a` and `b`.

b) Call the function with the last and the penultimate number of your students ID-Number. Print the result and than also your students ID-Number in following form - example: `My ID is: 1733456 and the result of myMultiplication is: 30` 
Explanation: ID = 1733456, last number = 6, penultimate = 5, result = 5\*6

*Hint:* there are different approaches to solving this. You can for example create a variable within the function to store the current sum and continually add the value of one variable; repeating the process as many times as indicated by the other variable. The `range()` function can be useful for executing a loop a certain number of times. 

<!-- grade_id: function, points: 0, grade: false, solution: true -->
```python
def myMultiplication(a = 5, b = 2):
    # YOUR CODE HERE
    raise NotImplementedError()

print(myMultiplication())
idnumber = myMultiplication(6, 5)
print(f"My ID is: 1733456 and the result of myMultiplication is: {idnumber}")
```

<!-- grade_id: function-test, points: 3, grade: true, solution: false -->
```python
assert myMultiplication(0,0) == 0
assert myMultiplication(1,0) == 0
assert myMultiplication(1,1) == 1
assert myMultiplication(6,5) == 30
```

## Exercise 1.5: If

Write a function called `isPositiveOrNegative()` that has a parameter `s`. The parameter should be able to accept a string containing a number written as text. 

Convert the string into a float (search for how to perform that; you will find a direct answer for example on Stack Overflow). Check if the input is greater, equal or smaller than zero. If it is greater than zero, return "positive"; if it is equal to zero return "zero"; and if it is smaller than zero return "negative".

<!-- grade_id: if, points: 0, grade: false, solution: true -->
```python
def isPositiveOrNegative(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# This should return negative
print(isPositiveOrNegative("-5.2"))
```

<!-- grade_id: if-test, points: 3, grade: true, solution: false -->
```python
assert isPositiveOrNegative("-5.2") == "negative"
assert isPositiveOrNegative("5") == "positive"
assert isPositiveOrNegative("0.1") == "positive"
assert isPositiveOrNegative("0") == "zero"
assert isPositiveOrNegative("-0.1") == "negative"
assert isPositiveOrNegative(0) == "zero"
assert isPositiveOrNegative(1.353) == "positive"
assert isPositiveOrNegative(-5) == "negative"
```